Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

If you find your kernel dead (e.g,, "In[*]" appears before a code cell), please interrupt the kernel (press "&#11035;").


Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE" and delete the statement "raise NotImplementedError()" 

Don't forget write down your name and collaborators below:

```python
NAME = "Zhenghao Wu"
ID = "1630003054"
```

---

# Python Functions

## Basic Features
- reusable chunks of code
- not run in a program until they are <b>"called"</b>
- has a name
- has <b>parameters</b>(0 or more) 
- has a <b>docstring</b>(optional but recommended)
- has a body
- <b>Returns</b> something

**We have introduced some basic knowledge about Python functions. Let us write a simple program to see if you have learnt how to define and call a function.**

**Task 1**
Enter a number, then the program should draw a square using *, which uses that number as the length. For example, ```draw_rect(4)``` will return a pattern as follows:
![](Gigi_img/dict_img6.png)

<!-- grade_id: cell-7c2303aa07eaff52, points: 3, grade: true, solution: true -->
```python
# Your code for task 1 here
def draw_rect(side_len):
    for i in range(side_len):
        for j in range(side_len):
            print("*    ", end="")
        print("\n")

n = int(input('Please input the side length(integer type) of the square: '))
draw_rect(n)
```

Output:
```
Please input the side length(integer type) of the square: 10
*    *    *    *    *    *    *    *    *    *    

*    *    *    *    *    *    *    *    *    *    

*    *    *    *    *    *    *    *    *    *    

*    *    *    *    *    *    *    *    *    *    

*    *    *    *    *    *    *    *    *    *    

*    *    *    *    *    *    *    *    *    *    

*    *    *    *    *    *    *    *    *    *    

*    *    *    *    *    *    *    *    *    *    

*    *    *    *    *    *    *    *    *    *    

*    *    *    *    *    *    *    *    *    *
```

## Function Arguments (Parameters)

- You can call a function by using the following types of formal arguments:
    - Required arguments
    - Keywords arguments
    - Default arguments
    - Variable-length argument
    

### Required arguments

Required arguments are the arguments passed to a function in correct positional order. Here, the number of arguments in the <i>function call</i> should match <b>exactly</b> with the <i>function definition</i>.

**Example**

```python
def printme(a,b):
    """This prints a passed string into this function"""
    print("a is ",a,"b is", b)
    
# Now you can call printme function
# Note, "1" and "2" are exactly matching to "a" and "b" here.
printme(1, 2)
```

Output:
```
a is  1 b is 2
```

### Keyword arguments
- Specify the parameter name when you call the function, .
- Arguments can be <b>out of order</b>.

<b>Example:</b>

```python
def printme(a,b):
    print("a is ", a, ", b is ", b)
    
printme(b = 1, a = 2)
```

Output:
```
a is  2 , b is  1
```

### Default arguments
- In the definition of the function, a default argument is set to the parameter.
- Setting the argument is optional when calling the function.
- The parameter without default value must be in front of the default argument.

Observe the following example:

```python
def printme(a, b = 10):
    print("a is: ", a, ", b is: ", b)
    
printme(1)
printme(a = 3)
printme(4, 5)
```

Output:
```
a is:  1 , b is:  10
a is:  3 , b is:  10
a is:  4 , b is:  5
```

### Variable-length arguments
- When calling the function, the number of the argument can vary.
- In the function definition, we use variable-length arguments.

Syntax for a function with non-keyword variable arguments is:
```python
def functionname([formal_args,] *var_args_tuple):
    function_suite
    return [expression]
```

   - An asterisk(*) is placed before the variable name that holds the values of all nonkeyword variable arguments. This tuple remains empty if no additional arguments are specified during the function call.

```python
def printinfo(a, *vartuple):
    """This prints a variable passed arguments"""
    print("Output is: ")
    print(a)
    for var in vartuple:
        print(var)
        
# sometimes, the number of values following "hi" is not fixed 
printinfo("hi", 1, 2, 3, 4)


```

Output:
```
Output is: 
hi
1
2
3
4
```

```python
another_name = printinfo
another_name('hi', 1, 2)
```

Output:
```
Output is: 
hi
1
2
```

```python
another_name == printinfo
```

Output:
```
True
```

## Return multiple values

A function can return only one value as you have played. In Python, returning multiple values can be done using a tuple.

```python
def test(x, y):
    a = x + y
    b = x - y
    # returning a and b, can also be written as return a, b
    return (a,b)

# calling the test function
(x, y) = test(2, 1)
# the result is 3, 1
print(x, y)
```

Output:
```
3 1
```

## Anonymous Functions and Lambda Operator
- These functions are called anonymous because they are not declared in the standard manner by using the ```def``` keyword. You can use the ```lambda``` keyword to create small anonymous functions.
    - Lambda forms can take any number of arguments but return just one value in the form of an expression. They cannot contain commands or multiple expression.
    - An anonymous function cannot be a direct call to print because lambda requires an expression.
    - Lambda functions have their own local namespace and cannot access variable other than those in their parameter list and those in their parameter list and those in the global namespace.
    
**Syntax**

```lambda [arg1 [, arg2, ...... argn]]: expression```

**Example**

```python
# Function definition is here
sum_lambda1 = lambda arg1, arg2: arg1 + arg2

print(sum_lambda1(1, 2))
```

Output:
```
3
```

```python
def sum_lambda2(arg1, arg2):
    return arg1 + arg2

sum_lambda2(1, 2)
```

Output:
```
3
```

```python
# Now you can call sum as a function
print("Value of total : ", sum_lambda1(10, 20))
print("Value of total : ", sum_lambda1(20, 20))
```

Output:
```
Value of total :  30
Value of total :  40
```

```python
(lambda arg1, arg2: arg1 + arg2)(10, 20)
```

Output:
```
30
```

## Pass by reference

- All parameters(arguments) in the Python language are passed by reference, which means if you change what a parameter refers to within a function, the change also reflects back in the calling function.

**Example**
- The example below shows the function maintaining the reference of the passed object and appending values in the same object.

```python
# Function definition is here
def changeme(mylist):
    # This changes a passed list into this function
    mylist.append([1, 2, 3, 4])
    print("Values inside the function:", mylist)
    
# Now you can call changeme function
mylist = [10, 20, 30]
changeme(mylist)
print("Values outside the function: ", mylist)
```

Output:
```
Values inside the function: [10, 20, 30, [1, 2, 3, 4]]
Values outside the function:  [10, 20, 30, [1, 2, 3, 4]]
```


## Scope of Variables
- Global variables
    - Defined outside the function
    - Can be accessed inside the functions
    - Using <b>global</b> keyword before modifying the value
- Local variables
    - Variables that are defined inside of a function
    - Objects outside the "scope" of the function will not be able to access that variable

```python
total = 0

def sum(arg1, arg2):
    # must use global keyword here
    global total
    # Since "total" is explicitly defined as globals, change of the variable will affect globally.
    total = arg1 + arg2;
    print("Inside the function tatal : ", total)
    
# Now you can call sum function
print("Initial total: ", total)
sum(10, 20)
print("Outside the function total : ", total)
```

Output:
```
Initial total:  0
Inside the function tatal :  30
Outside the function total :  30
```

**Task 2:**

- Write a program that asks the user to enter two points on a 2D plane (i.e. enter X1 & Y1, enter X2 & Y2). Compute the distance between those points using functions.
- Continually ask the user for numbers until they wish to quit

<!-- grade_id: cell-71f654aca03f6f71, points: 0, grade: false, solution: true -->
```python

# Tips: define a function cmpt_distance(p1,p2) to compute the distance between p1 and p2
# p1 = (x1,y1) and p2 = (x2,y2)


import math
def cmpt_distance(p1,p2):
    return math.sqrt((int(p1[0])-int(p2[0]))**2+(int(p1[1])-int(p2[1]))**2)
```

<!-- grade_id: cell-e156f40c1e1dadd1, points: 2, grade: true, solution: false -->
```python
p1 = (1, 2)
p2 = (4, 6)
assert cmpt_distance(p1, p2) == 5
```

<!-- grade_id: cell-6de37cb1375fa559, points: 3, grade: true, solution: true -->
```python
# the main function to realize that ask the user for numbers until they wish to quit.
# eop is the end sign of the program
eop = int(input('Input a natural number (int 0 to exit) '))

while (eop is not 0):
    p1 = input("please input x1, and y1").split(" ")
    p2 = input("please input x2, and y2").split(" ")
    print("Two point distance: %f" % cmpt_distance(p1, p2))
    eop = int(input('Input a natural number (int 0 to exit) '))

print("Quit!")
```

Output:
```
Input a natural number (int 0 to exit) 1
please input x1, and y11 2
please input x2, and y24 6
Two point distance: 5.000000
Input a natural number (int 0 to exit) 0
Quit!
```

**Task 3:**

- Define a function ```decimal2Bin``` to return a string, which is the binary form(str type) for a natural number.

<!-- grade_id: cell-28eaf24e88fd65ef, points: 0, grade: false, solution: true -->
```python

def decimal2Bin(inDecimal):
    if inDecimal == 0 or inDecimal == 1:
        return inDecimal
    return str(decimal2Bin(int(inDecimal / 2))) + str(inDecimal % 2)
```

<!-- grade_id: cell-3dd45a7dee1aaeb4, points: 4, grade: true, solution: false -->
```python
assert decimal2Bin(14) == '1110'
```

**Task 4:**

- Define a function ```find_RepStr``` to return a list, which includes all the consecutive characters which are also sub-strings appeared in the object string.

<!-- grade_id: cell-8ed03d64c71d9334, points: 0, grade: false, solution: true -->
```python
# using dir(str) to find the built-in function of python object, str, 
# and find the way to solve the problem
# YOUR CODE HERE

def find_RepStr(inputStr):
    RepStrArray = []
    for i in range(len(inputStr)):
        if i == len(inputStr) - 1:
            break
        else:
            indexEnd = i + 1
        while inputStr[i] == inputStr[indexEnd]:
            if(indexEnd >= len(inputStr) - 1):
                break
            indexEnd += 1
        if indexEnd - i > 1:
            RepStrArray.append(inputStr[i:indexEnd])
    return RepStrArray
```

<!-- grade_id: cell-dd569bd0042fefab, points: 4, grade: true, solution: false -->
```python
assert find_RepStr("http://www.google.com") == ['tt', '//', 'www', 'ww', 'oo']
assert find_RepStr("abcdefgfedcba") == []
assert find_RepStr("aaaabbbccd") == ['aaaa', 'aaa', 'aa', 'bbb', 'bb', 'cc']
```

**Task 5:**

- Define a function ```cmpt_ticket``` to compute the admission price of a park. The detailed rules are following:
    - admission price on weekends is 1.2 times of that on weekdays
    - admission price for a child is a half of an adult's 

<!-- grade_id: cell-86a4e1aeebbf031d, points: 0, grade: false, solution: true -->
```python

def cmpt_ticket(price, child = False, weekend = False):
    if child:
        price /= 2
    if weekend:
        price *= 1.2
    return price

```

<!-- grade_id: cell-d3a46b83d3601dc8, points: 4, grade: true, solution: false -->
```python
assert cmpt_ticket(100) == 100
assert cmpt_ticket(100, child = True) == 50
assert cmpt_ticket(100, weekend = True) == 120
assert cmpt_ticket(100, True, True) == 60
```

```python

```